Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nahom Agize"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq




def add_to_median_heap(minh, maxh, elem):
    if maxh == []:
        heapq.heappush(maxh, elem)
    if minh == [] and maxh != []:
        heapq.heappush(minh, elem)

    if elem < maxh[0]:
        heapq.heappush(maxh, elem)
        heapq._heapify_max(maxh)

    if elem > minh[0]:
        heapq.heappush(minh, elem)
        heapq.heapify(minh)

        
        
    
    #raise NotImplementedError()
```

## Question 2 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    if len(minh) > len(maxh):
        median = minh[len(minh)-1]
    elif len(minh) < len(maxh):
        median = maxh[len(maxh)-1]
    else:
        median = (len(maxh)+len(minh))//2
    return median

    raise NotImplementedError()
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []

for a in range(1,100, 2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
3
5
7
9
11
13
15
17
19
21
23
25
27
29
31
33
35
37
39
41
43
45
47
49
51
53
55
57
59
61
63
65
67
69
71
73
75
77
79
81
83
85
87
89
91
93
95
97
99
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The worst case complexity for this function would be O(logN)

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
Finding the median basically included comparing and selecting the element to be popped off, thus will have worst case complexity of O(1)

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
This is much more efficient compared to the "vanilla" way. This is because we do not really need an already sorted array to find the median, and if sorted will take O(nlogN) while this algorithm computes the median in O(logN) time. 

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def selectsort(lst): #modified quick sort
    #select the median of the first, mid, and last element of the list as pivot
    locations = [0.0, 0.5, 1.0 - 1e-16]
    N = len(lst)
    inds = [lst[int(N * n)] for n in locations]
    values = [ind for ind in inds]
    if (values[0]<values[1]<values[2]) or (values[2]<values[1]<values[0]):
        partition = values[1]
    elif (values[0]<values[2]<values[1]) or (values[1]<values[2]<values[0]):
        partition = values[2]
    else:
        partition = values[0]
    
    #divide the list into three parts
    lower = [a for a in lst if a < partition]
    upper = [a for a in lst if a > partition]
    mid = [a for a in lst if a == partition]
    
    return (lower,mid,upper)
    
    
#k here indicates the k-th element with 1-based index
def qselect(lst,k):
    #divide the list
    sort = selectsort(lst)
    left = sort[0]
    mid = sort[1]
    right = sort[2]
    #if and only if the lower part has k-1 elements, the pivot is the k-th smallest element
    #if the lower part has more than k-1 elements, we recursively select the lower part only
    #if the lower part has less than k-1 elements, we recursively select the upper part with k-len(left)
    if len(left)>k-1:
        return qselect(left,k)
    #if the lower part has no elements - this only occurs when there are only 2 elements in the list (so that one
    #goes to mid and the other goes to upper), we simply choose the elements based on the value of k (because k
    #can only be either 1 or 2)
    if len(left)==0:
        if k==2:
            return right[-1]
        if k==1:
            return mid[0]
    if len(left)<k-1:
        return qselect(mid+right,k-len(left))
    else:
        return mid[0]

     

    

```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

Output:
```
<ipython-input-227-abfbba6b32f6> in selectsort(lst)
      3     locations = [0.0, 0.5, 1.0 - 1e-16]
      4     N = len(lst)
----> 5     inds = [lst[int(N * n)] for n in locations]
      6     values = [ind for ind in inds]
      7     if (values[0]<values[1]<values[2]) or (values[2]<values[1]<values[0]):

<ipython-input-227-abfbba6b32f6> in <listcomp>(.0)
      3     locations = [0.0, 0.5, 1.0 - 1e-16]
      4     N = len(lst)
----> 5     inds = [lst[int(N * n)] for n in locations]
      6     values = [ind for ind in inds]
      7     if (values[0]<values[1]<values[2]) or (values[2]<values[1]<values[0]):

IndexError: list index out of range
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
random.seed(123)
lst = [a for a in range(10)]
random.shuffle(lst)
for a in range(1,1001): #because the k is 1-based, I need to change the value of a here
    print(qselect(lst, a))
```

Output:
```
0
1
2
3
4
5
6
7
8
9
... last 1 frames repeated, from the frame below ...

<ipython-input-222-c8910bcaa39a> in qselect(lst, k)
     41             return mid[0]
     42     if len(left)<k-1:
---> 43         return qselect(mid+right,k-len(left))
     44     else:
     45         return mid[0]

RecursionError: maximum recursion depth exceeded while calling a Python object
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
Each time we don't need to actually sort the list, we only need to partition it. The number of steps is relevant to the number of comparisons. For the average case, we expect to divide the list by half-half. Since we only need to recursively solve one half of the list, the recurrence relation would be n+n/2+4/n+8/n+....<2n.

## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
To solve for the recurrence time for the best case, that is, we only need to do one partition. The time would be n.

## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
For the worst case, it is similar to that of the quick sort. Each time we divide the list into 1 and n-1 in the end, it would take n^2 time in the worst scenario.

```python

```